## Convert to Binary

<!-- grade_id: d39095, points: 0, grade: false, solution: true -->
```python
def bin(n):
    """
    Convert n from decimal to binary. Return the binary digits in a list in the order from highest place value to lowest place value.
    """
    result = []
    n = 10
    q = n
    while(q>0):
        r = q %2
        q = q//2
        result.append(r)
    for i in result:
        return(result[i])
    raise NotImplementedError()
```

## Testing

<!-- grade_id: 7dd5ce, points: 1, grade: true, solution: false -->
```python
# Check that bin function returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(bin(34), [1, 0, 0, 0, 1, 0])
assert_equal(bin(42), [1, 0, 1, 0, 1, 0])
assert_equal(bin(121), [1, 1, 1, 1, 0, 0, 1])


```

Output:
```
/usr/lib/python3.8/unittest/case.py in assertEqual(self, first, second, msg)
    910         """
    911         assertion_func = self._getAssertEqualityFunc(first, second)
--> 912         assertion_func(first, second, msg=msg)
    913 
    914     def assertNotEqual(self, first, second, msg=None):

/usr/lib/python3.8/unittest/case.py in _baseAssertEqual(self, first, second, msg)
    903             standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    904             msg = self._formatMessage(msg, standardMsg)
--> 905             raise self.failureException(msg)
    906 
    907     def assertEqual(self, first, second, msg=None):

AssertionError: '0b100010' != [1, 0, 0, 0, 1, 0]
```

```python

```